Genetic and physical interactions reveal overlapping and distinct contributions to meiotic double-strand break formation in C. elegans

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The manuscript by Raices et al., provides some novel insights into the role and interactions between SPO-11 accessory proteins in C. elegans. The authors propose a model of meiotic DSBs regulation, critical to our understanding of DSB formation and ultimately crossover regulation and accurate chromosome segregation. The work also emphasizes the commonalities and species-specific aspects of DSB regulation.

Strengths:

This study capitalizes on the strengths of the C. elegans system to uncover genetic interactions between a lSPO-11 accessory proteins. In combination with physical interactions, the authors synthesize their findings into a model, which will serve as the basis for future work, to determine mechanisms of DSB regulation.

Weaknesses:

The methodology, although standard, still lacks some rigor, especially with the IPs.

Reviewer #2 (Public review):

Summary:

Meiotic recombination initiates with the formation of DNA double-strand break (DSB) formation, catalyzed by the conserved topoisomerase-like enzyme Spo11. Spo11 requires accessory factors that are poorly conserved across eukaryotes. Previous genetic studies have identified several proteins required for DSB formation in C. elegans to varying degrees; however, how these proteins interact with each other to recruit the DSB-forming machinery to chromosome axes remains unclear.

In this study, Raices et al. characterized the biochemical and genetic interactions among proteins that are known to promote DSB formation during C. elegans meiosis. The authors examined pairwise interactions using yeast two-hybrid (Y2H) and co-immunoprecipitation and revealed an interaction between a chromatin-associated protein HIM-17 and a transcription factor XND-1. They further confirmed the previously known interaction between DSB-1 and SPO-11 and showed that DSB-1 also interacts with a nematodespecific HIM-5, which is essential for DSB formation on the X chromosome. They also assessed genetic interactions among these proteins, categorizing them into four epistasis groups by comparing phenotypes in double vs. single mutants. Combining these results, the authors proposed a model of how these proteins interact with chromatin loops and are recruited to chromosome axes, offering insights into the process in C. elegans compared to other organisms.

Weaknesses:

This work relies heavily on Y2H, which is notorious for having high rates of false positives and false negatives. Although the interactions between HIM-17 and XND-1 and between DSB-1 and HIM-5 were validated by co-IP, the significance of these interactions was not tested in vivo. Cataloging Y2H and genetic interactions does not yield much more insight. The model proposed in Figure 4 is also highly speculative.

Reviewer #3 (Public review):

The goal of this work is to understand the regulation of double-strand break formation during meiosis in C. elegans. The authors have analyzed physical and genetic interactions among a subset of factors that have been previously implicated in DSB formation or the number of timing of DSBs: CEP-1, DSB-1, DSB-2, DSB-3, HIM-5, HIM-17, MRE-11, REC-1, PARG-1, and XND-1.

The 10 proteins that are analyzed here include a diverse set of factors with different functions, based on prior analyses in many published studies. The term "Spo11 accessory factors" has been used in the meiosis literature to describe proteins that directly promote Spo11 cleavage activity, rather than factors that are important for the expression of meiotic proteins or that influence the genome-wide distribution or timing of DSBs. Based on this definition, the known SPO-11 accessory factors in C. elegans include DSB-1, DSB2, DSB-3, and the MRN complex (at least MRE-11 and RAD-50). These are all homologs of proteins that have been studied biochemically and structurally in other organisms. DSB-1 & DSB-2 are homologs of Rec114, while DSB-3 is a homolog of Mei4. Biochemical and structural studies have shown that Rec114 and Mei4 directly modulate Spo11 activity by recruiting Spo11 to chromatin and promoting its dimerization, which is essential for cleavage. The other factors analyzed in this study affect the timing, distribution, or number of RAD-51 foci, but they likely do so indirectly. As elaborated below, XND-1 and HIM-17 are transcription factors that modulate the expression of other meiotic genes, and their role in DSB formation is parsimoniously explained by this regulatory activity. The roles of HIM-5 and REC-1 remain unclear; the reported localization of HIM-5 to autosomes is consistent with a role in transcription (the autosomes are transcriptionally active in the germline, while the X chromosome is largely silent), but its loss-of-function phenotypes are much more limited than those of HIM-17 and XND-1, so it may play a more direct role in DSB formation. The roles of CEP-1 (a Rad53 homolog) and PARG-1 are also ambiguous, but their homologs in other organisms contribute to DNA repair rather than DSB formation.

We appreciate the reviewer’s clarification. However, the definition of Spo11 accessory factors varies across the literature. Only Keeney and colleagues define these as proteins that physically associate with and activate Spo11 to catalyze DSB formation (Keeney, Lange & Mohibullah, 2014; Lam & Keeney, 2015). In contrast, other authors have used the term more broadly to refer to proteins that promote or regulate Spo11-dependent DSB formation, without necessarily implying a direct interaction with Spo11 (e.g., Panizza et al., 2011; Robert et al., 2016; Stanzione et al., 2016; Li et al., 2021; Lange et al., 2016). Thus, our usage of the term follows this broader functional definition.

An additional significant limitation of the study, as stated in my initial review, is that much of the analysis here relies on cytological visualization of RAD-51 foci as a proxy for DSBs. RAD-51 associates transiently with DSB sites as they undergo repair and is thus limited in its ability to reveal details about the timing or abundance of DSBs since its loading and removal involve additional steps that may be influenced by the factors being analyzed.

We agree with the reviewer that counting RAD-51 foci provides only an indirect measure of SPO-11–dependent DSBs, as RAD-51 marks sites of repair rather than the breaks themselves. However, we would like to clarify that our current study does not rely on RAD51 foci quantification for any of the analyses or conclusions presented. None of the figures or datasets in this manuscript are based on RAD-51 cytology. Instead, our conclusions are drawn from genetic interactions, biochemical assays, and protein–protein interaction analyses.

The paper focuses extensively on HIM-5, which was previously shown through genetic and cytological analysis to be important for breaks on the X chromosome. The revised manuscript still claims that "HIM-5 mediates interactions with the different accessory factors sub-groups, providing insights into how components on the DNA loops may interact with the chromosome axis." The weak interactions between HIM-5 and DSB-1/2 detected in the Y2H assay do not convincingly support such a role. The idea that HIM-5 directly promotes break formation is also inconsistent with genetic data showing that him5 mutants lack breaks on the X chromosomes, while HIM-5 has been shown to be is enriched on autosomes. Additionally, as noted in my comment to the authors, the localization data for HIM-5 shown in this paper are discordant with prior studies; this discrepancy should be addressed experimentally.

We appreciate the reviewer’s concerns regarding the interpretation of HIM-5 function. The weak Y2H interactions between HIM-5 and DSB-1 are not interpreted as direct biochemical evidence of a strong physical interaction, but rather as a potential point of regulatory connection between these pathways. Importantly, these Y2H data are further supported by co-immunoprecipitation experiments, genetic interactions, and the observed mislocalization of HIM-5 in the absence of DSB-1. Together, these complementary results strengthen our conclusion that HIM-5 functionally associates with DSB-promoting complexes.

Regarding HIM-5 localization, the pattern we observe using both anti-GFP staining of the eaIs4 transgene (Phim-5::him-5::GFP) and anti-HA staining of the HIM-5::HA strain is consistent with that reported by McClendon et al. (2016), who validated the same eaIs4 transgene. Although the pattern difers slightly from Meneely et al. (2012), that used a HIM5 antibody that is no longer functional and that has been discontinued by the commercial source. In this prior study, a weak signal was detected in the mitotic region and late pachytene, but stronger signal was seen in early to mid-pachytene. Our imaging— optimized for low background and stable signal—similarly shows robust HIM-5 localization in early and mid-pachytene, supporting the reliability of our GFP and HA-tagged analyses.

The recent analysis of DSB formation in C. elegans males (Engebrecht et al; PloS Genetics; PMID: 41124211) shows that in absence of him-5 there is a significant reduction of CO designation (measured as COSA-1 foci) on autosomes. This study strongly supports a direct and general role for HIM-5 in crossover formation— on both autosomes and on the hermaphrodite X.

This paper describes REC-1 and HIM-5 as paralogs, based on prior analysis in a paper that included some of the same authors (Chung et al., 2015; DOI 10.1101/gad.266056.115). In my initial review I mentioned that this earlier conclusion was likely incorrect and should not be propagated uncritically here. Since the authors have rebutted this comment rather than amending it, I feel it is important to explain my concerns about the conclusions of previous study. Chung et al. found a small region of potential homology between the C. elegans rec-1 and him-5 genes and also reported that him-5; rec-1 double mutants have more severe defects than either single mutant, indicative of a stronger reduction in DSBs. Based on these observations and an additional argument based on microsynteny, they concluded that these two genes arose through recent duplication and divergence. However, as they noted, genes resembling rec-1 are absent from all other Caenorhabditis species, even those most closely related to C. elegans. The hypothesis that two genes are paralogs that arose through duplication and divergence is thus based on their presence in a single species, in the absence of extensive homology or evidence for conserved molecular function. Further, the hypothesis that gene duplication and divergence has given rise to two paralogs that share no evident structural similarity or common interaction partners in the few million years since C. elegans diverged from its closest known relatives is implausible. In contrast, DSB-1 and DSB-2 are both homologs of Rec114 that clearly arose through duplication and divergence within the Caenorhabditis lineage, but much earlier than the proposed split between REC-1 and HIM-5. Two genes that can be unambiguously identified as dsb-1 and dsb-2 are present in genomes throughout the Elegans supergroup and absent in the Angaria supergroup, placing the duplication event at around 18-30 MYA, yet DSB-1 and DSB-2 share much greater similarity in their amino acid sequence, predicted structure, and function than HIM-5 and REC-1. Further, Raices place HIM-5 and REC-1 in different functional complexes (Figure 3B).

We respectfully disagree with the reviewer’s characterization of the relationship between HIM-5 and REC-1. Our use of the term “paralog” follows the conclusions of Chung et al. (2015), a peer-reviewed study that provided both sequence and microsynteny evidence supporting this relationship. While we acknowledge that the degree of sequence conservation is limited, the evolutionary scenario proposed by Chung et al. remains the only published framework addressing this question. Further the degree of homology between either HIM-5 or REC-1 and the ancestral locus are similar to that observed for DSB-1 and DSB-2 with REC-114 (Hinman et al., 2021). We therefore retain the use of the term “paralog” in reference to these genes. Importantly, our conclusions regarding their distinct molecular and functional roles are independent of this classification.

The authors acknowledge that HIM-17 is a transcription factor that regulates many meiotic genes. Like HIM-17, XND-1 is cytologically enriched along the autosomes in germline nuclei, suggestive of a role in transcription. The Reinke lab performed ChIP-seq in a strain expressing an XND-1::GFP fusion protein and showed that it binds to promoter regions, many of which overlap with the HIM-17-regulated promoters characterized by the Ahringer lab (doi: 10.1126/sciadv.abo4082). Work from the Yanowitz lab has shown that XND-1 influences the transcription of many other genes involved in meiosis (doi: 10.1534/g3.116.035725) and work from the Colaiacovo lab has shown that XND-1 regulates the expression of CRA-1 (doi: 10.1371/journal.pgen.1005029). Additionally, loss of HIM-17 or XND-1 causes pleiotropic phenotypes, consistent with a broad role in gene regulation. Collectively, these data indicate that XND-1 and HIM-17 are transcription factors that are important for the proper expression of many germline-expressed genes. Thus, as stated above, the roles of HIM-17 and XND-1 in DSB formation, as well as their effects on histone modification, are parsimoniously explained by their regulation of the expression of factors that contribute more directly to DSB formation and chromatin modification. I feel strongly that transcription factors should not be described as "SPO-11 accessory factors."

The ChIP analysis of XND-1 binding sites (using the XND-1::GFP transgene we provided to the Reinke lab) was performed, and Table S3 in the Ahringer paper suggests it is found at germline promoters, although the analysis is not actually provided. We completely agree that at least a subset of XND-1 functions is explained by its regulation of transcriptional targets (as we previously showed for HIM-5). However, like the MES proteins, a subset of which are also autosomal and impact X chromosome gene expression, XND-1 could also be directly regulating chromatin architecture which could have profound effects on DSB formation. As stated in our prior comments, precedent for the involvement of a chromatin factor in DSB formation is provided by yeast Spp1.

Recommendations for the authors:

Editor comments:

As you can see, the reviewers have additional comments, and the authors can include revisions to address those points prior to publicizing 'a version of record' (e.g. hatching rate assay mentioned by reviewer #1). This type of study, trying to catalog interactions of many factors, inevitably has loose ends, but in my opinion, it does not reduce the value of the study, as long as statements are not misleading. I suggest that the authors address issues by making changes to the main text. After the next round of adjustments by authors, I feel that it will be ready for a version of record, based on the spirit of the current eLife publication model.

Reviewer #1 (Recommendations for the authors):

I still have concerns about the HIM-17 IP and immunoblot probing with XND-1 antibodies. While the newly provided whole extract immunoblot clearly shows a XND-1 specific band that goes away in the mutant extracts, there is additional bands that are recognized - the pattern looks different than in the input in Figure 1B. Additionally, there is still a band of the corresponding size in the IPs from extracts not containing the tagged allele of HIM-17, calling into question whether XND-1 is specifically pulled down.

The authors did not include the hatching rate as pointed out in the original reviews. In the rebuttal:

"Great question. I guess we need to do this while back out for review. If anyone has suggestions of what to say here. Clearly we overlooked this point but do have the strain."

We thank the reviewer for this suggestion. We had intended to include a hatching analysis; however, during the course of this work we discovered that our him-17 stock had acquired an additional linked mutation(s) that altered its phenotype and led to inconsistent results. This strain was used to rederive the him-17; eaIs4 double mutant after our original did not survive freeze/thaw. Given the abnormal behavior observed in this line, we concluded that proceeding with the hatching assays could yield unreliable data. We are currently reestablishing a verified him-17 strain, but in the interest of accuracy and reproducibility, we have restricted our analysis in this manuscript to validated datasets derived from confirmed strains.

Reviewer #2 (Recommendations for the authors):

The authors have addressed most of the previous concerns and substantially improved the manuscript. The new data demonstrate that HIM-5 localization depends on DSB-1, and together with the Y2H and Co-PI results, strengthen the link between HIM-5 and the DSBforming machinery in C. elegans. The remaining points are outlined below:

Specific comments:

The font size of texts and labels in the Figure is very small and is hardly legible. Please enlarge them and make them clearly visible (Fig 1A, 1B, 2A, 2B, 2C, 2D, 2E, 3A, 3B, 3C, 3D, 3F)

Done

Although the authors have addressed the specificity of the XND-1 antibody, it remains unclear whether the boxed band is specific to the him-17::3xHA IP, since the same band appears in the control IP, albeit with lower intensity (Fig 1B). Is the ~100 kDa band in the him-17::3xHA IP a modified form XND-1? While antibody specificity was previously demonstrated by IF using xnd-1 mutants, it would be ideal to confirm this on a western blot as well.

A Western Blot performed using whole cell extracts and probed with the anti- XND-1 antibody has been provided in the revised version of the manuscript (Fig. S1A). This confirms that the antibody specifically recognizes XND-1 protein. We believe that the ~100 kDa band mentioned by the reviewer is likely to be a non-specific cross reaction band detected by the antibody, since an identical band of the same mW was also detected in xnd-1 null mutants (Fig. S1A).

Regarding the IP negative controls, we are firmly convinced the boxed band to be specific, and the fact that a (very) low intensity band is also found in the negative control should not infringe the validity of the HIM-17-XND-1 specific interaction. There is a constellation of similar examples present across the literature, as it is widely acknowledged amongst biochemists that some proteins may “stick” to the beads due their intrinsic biochemical properties despite usage of highly stringent IP buffers. However, the high level of enrichment detected in the IP (as also underlined by the reviewer) corroborates that XND-1 specifically immunoprecipitates with HIM-17 despite a low, non-specific binding to the HA beads is present. If interaction between XND-1 and HIM-17 was non-specific, we logically would have found the band in the IP and the band in the negative control to be of very similar intensity, which is clearly not the case.

Although co-IP assays are generally considered not a strictly quantitative assay, we want to emphasize that a comparable amount of nuclear extract was employed in both samples as also evidenced by the inputs, in which it is also possible to see that if anything, slightly less nuclear extracts were employed in the him-17::3xHA; him-5::GFP::3xFLAG vs. the him5::GFP::3xFLAG negative control, corroborating the above mentioned points.

Lastly, it is crucial to mention that mass spectrometry analyses performed on HIM17::3xHA pulldowns show XND-1 as a highly enriched interacting protein (Blazickova et al.; 2025 Nature Comms.), which strongly supports our co-IP results.

The subheading "HIM-5 is the essential factor for meiotic breaks in the X chromosome" does not accurately represent the work described in the Results or in Figure 1. I disagree with the authors' response to the earlier criticism. The issue is not merely semantic. The data do not demonstrate that HIM-5 is required for DSB formation on the X chromosome - this conclusion can only be inferred. What Figure 1 shows is that XND-1 and HIM-17 interact, and that pie-1p-driven HIM-5 expression can partially rescue meiotic defects of him-17 mutants. This supports the conclusion that him-5 is a target of HIM-17/XND-1 in promoting CO formation on the X chromosome. However, the data provide no direct evidence for the claim stated in the subheading. I strongly encourage authors to revise the subheading to more accurately represent the findings presented in the paper.

After considering the reviewer’s comments, we have revised the subheading to more accurately describe our findings.

In Fig1C, please fix the typo in the last row - "pie1p::him5-::GFP" to "pie-1p::him- 5::GFP".

Done

In Fig 2C, "p" is missing from the label on the right for Phim-5::him-5::GFP.

Done

In Fig 3I, bring the labels (DSB-1/2/3) at the lower right to the front.

Done

In Concluding Remarks, please fix the typo "frequently".

Done

Reviewer #3 (Recommendations for the authors):

The experiments that analyze HIM-5 in dsb-1 mutants should be repeated using antibodies against the endogenous HIM-5 antibody, and localization of the HIM-5::HA and HIM-5::GFP proteins should be compared directly to antibody staining. This work uses an epitopetagged protein and a GFP-tagged protein to analyze the localization of HIM-5, while prior work (Meneely et al., 2012) used an antibody against the endogenous protein. In Figures 2 and S4 of this paper, neither HIM-5::HA nor HIM-5::GFP appears to localize strongly to chromatin, and autosomal enrichment of HIM-5, as previously reported for the endogenous protein based on antibody staining, is not evident. Moreover, HIM-5::GFP and HIM-5::HA look different from each other, and neither resembles the low-resolution images shown in Figure 6 in Meneely et al 2012, which showed nuclear staining throughout the germline, including in the mitotic zone, and also in somatic sheath cells. Given the differences in localization between the tagged transgenes and the endogenous protein, it is important to analyze the behavior of the endogenous, untagged protein. A minor issue: a wild-type control should also be shown for HIM-5::HA in Figure S4.

Wild type control added to figure S4

Evidence that XND-1 and HIM-17 form a complex is weak; it is supported by the Y2H and co-IP data but opposed by functional analysis or localization. The diversity of proteins found in the Co-IP of HIM-17::GFP (Table S2) indicate that these interactions are unlikely to be specific. The independent localization of these proteins to chromatin is clear evidence that they do not form an obligate complex; additionally, they have been found to regulate distinct (although overlapping) sets of genes. The predicted structure generated by Alphafold3 has very low confidence and should not be taken as evidence for an interaction.The newly added argument about the lack of apparently overlap between HIM-17 and XND1 due to the distance between the HA tag on HIM-17 and XND-1 is flawed and should be removed - the extended C-terminus in the predicted AlphaFold3 C-terminus of HIM-17 has been interpreted as if it were a structured domain. Moreover, the predicted distance of 180 Å (18 nm) is comparable to the distance between a fluorophore on a secondary antibody and the epitope recognized by the primary antibody (~20-25 nm) and is far below than the resolution limit of light microscopy.

We appreciate the reviewer’s thoughtful comment. The evidence supporting a physical interaction between XND-1 and HIM-17 is not only shown by our co-IP experiments, but it has also been recently shown in an independent study where MS analyses were conducted on HIM-17::3xHA pull downs to identify novel HIM-17 interactors (Blazickova et al.; 2025 Nature Comms). As shown in the data provided in this study, also under these experimental settings XND-1 was identified as a highly enriched putative HIM-17 interactor. We do acknowledge that their chromatin localization patterns are distinct and they regulate overlapping but not identical sets of genes, however, it is worth noting that protein–protein interactions in meiosis are often transient or context-dependent, and may not necessarily result in co-localization detectable by microscopy. In line with this, in the same work cited above, a similar situation for BRA-2 and HIM-17 was reported, as they were shown to interact biochemically despite the absence of overlapping staining patterns.

Minor issues:

The images shown in Panel D in Figure 1 seem to have very different resolutions; the HTP3/HIM-17 colocalization image is particularly blurry/low-resolution and should be replaced. The contrast between blue and green cannot be seen clearly; colors with stronger contrast should be used, and grayscale images should also be shown for individual channels. High-resolution images should probably be included for all of the factors analyzed here to facilitate comparisons.

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Double-strand breaks (DSBs) are the most deleterious lesions experienced by our genome. Yet, DSBs are intentionally induced during gamete formation to promote the exchange of genetic material between homologous chromosomes. While the conserved topoisomerase-like enzyme Spo11 catalyzes DSBs, additional regulatory proteins— referred to as “Spo11 accessory factors”— regulate the number, timing, and placement of DSBs during meiotic prophase ensuring that SPO11 does not wreak havoc on the genome. Despite the importance of the accessory factors, they are poorly conserved at the sequence level suggesting that these factors may adopt unique functions in different species. In this work, we present a detailed analysis of the genetic and physical interactions between the DSB factors in the nematode Caenorhabditis elegans providing new insights into conserved and novel functions of these proteins. This work shows that HIM-5 is the determinant of X-chromosome-specific crossovers and that its retention in the nucleus is dependent on DSB-1, the sole accessory factor that interacts with SPO-11. We further provide evidence that HIM-5 mediates interactions with the different accessory factors sub-groups, providing insights into how components on the DNA loops may interact with the chromosome axis.

Introduction

One of the seminal events during meiosis is the formation of crossovers (CO) during the first meiotic division. While COs create new combinations of alleles through the exchange of genetic material between homologs, they also establish physical connections that ensure their proper alignment on the meiotic spindle and their subsequent segregation to opposite ends of the spindle. In the absence of COs, homologs segregate randomly, leading to gametes with altered chromosome numbers and subsequent fetal aneuploidies. In humans, aneuploidy is observed in 1 of 160 live births (Dungan, 2002) and in >50% of miscarriages (Hassold et al., 2007) underscoring the importance of proper CO formation.

A necessary early step in CO formation is the creation of DNA double-strand breaks (DSBs) catalyzed by the widely conserved Spo11 enzyme (Bergerat et al., 1997; Keeney et al., 1997). Although Spo11 is essential for DSB formation, it does not function alone: additional proteins regulate the timing, placement, and number of DSBs (Cole et al., 2010; Hinman et al., 2021; Keeney, 2008, Vrielynck et al., 2021). In Saccharomyces cerevisiae, where DSB formation has been best characterized, at least nine other proteins interact with Spo11 to regulate its recruitment and activation (Lam & Keeney, 2015; Martini et al., 2006). In vivo and in vitro physical interaction studies have revealed that these proteins form several subcomplexes. Briefly, the RMM complex (Rec114, Mer2, and Mei4) is loaded on chromosomes early in prophase and is required for Spo11 binding to sites of DNA cleavage (J. Li et al., 2006; Sasanuma et al., 2007). Activation of Mer2 requires phosphorylation by cyclin-dependent kinase (CDK) (Henderson et al., 2006), suggesting that Mer2 links DSB formation to the progression of the meiotic program (Arora et al., 2004; Henderson et al., 2006; J. Li et al., 2006; Maleki et al., 2007). Rec102 and Rec104 form another subcomplex that facilitates Spo11 dimerization and its association with DSB sites (Arora et al., 2004; Kee & Keeney, 2002; Prieler et al., 2005; Sasanuma et al., 2007). Ski8 binds Spo11 and is important for Spo11 nuclear localization (Arora et al., 2004). Finally, the Mre11–Rad50–Xrs2 (MRX) complex is also required for DSB formation but requires all the other DSB proteins to associate with DSB sites (Bouuaert et al., 2021). Despite their importance, most of these proteins are poorly conserved at the amino acid level although homologs have been identified for several of these in organisms ranging from plants to worms to mice (Cole et al., 2010; Kumar et al., 2010b). Importantly, a recent study of plants has presented a similar model for DSB induction as in yeast, yet the specific players and details differ (Vrielynck et al., 2021) highlighting the value of comparative studies to understand the evolutionary history of the DSB machinery in metazoans.

Earlier version: In vitro biochemical studies have shown that C. elegans SPO-11 is monomeric but that even when it is able to bind dsDNA, it does not exhibit DNA cleavage activity (Yeh et al., 2017). This observation strongly supports the hypothesis that cofactors are needed for DSB induction. DSB-1 and DSB-2, are homologs of the conserved Rec114 protein (Tesse et al, 2017), and DSB-3, a homolog of Mei4 (Hinman et al., 2021). These factors localize to chromosomes during early meiotic prophase and promote the activity of the SPO-11 complex (Rosu et al., 2013; Stamper et al, 2013; Hinman et al, 2021), functioning analogously to their counterparts in other species (Kumar et al., 2010a, 2015, 2018; Li et al, 2006; Maleki et al., 2007). him-17, mre-11, and rad-50 are also essential for the introduction of breaks in this organism (Chin & Villeneuve, 2001; Hayashi et al., 2007; Hinman et al., 2021; Reddy & Villeneuve, 2004; Stamper et al., 2013).

Revised version: In vitro biochemical studies have shown that C. elegans SPO-11 is monomeric but that even when it is able to bind dsDNA, it does not exhibit DNA cleavage activity (Yeh et al., 2017). This observation strongly supports the hypothesis that cofactors are needed for DSB induction. DSB-1 and DSB-2, are homologs of the conserved Rec114 protein (Tesse et al, 2017), and DSB-3, a homolog of Mei4 (Hinman et al., 2021). These factors localize to chromosomes during early meiotic prophase and promote the activity of the SPO-11 complex (Rosu et al., 2013; Stamper et al, 2013; Hinman et al, 2021), functioning analogously to their counterparts in other species (Kumar et al., 2010a, 2015, 2018; Li et al, 2006; Maleki, 2007). him-17, mre-11, and rad-50 are also essential for the introduction of breaks in this organism (Chin & Villeneuve, 2001; Hayashi et al., 2007; Hinman et al., 2021; Reddy & Villeneuve, 2004; Stamper et al., 2013).

Other genes, including xnd-1, him-5, rec-1, cep-1, and parg-1 have more limited roles in break induction. him-5 is fascinating in that it is required to ensure the complement of DSBs on the X chromosome (Janisiw et al., 2020; Meneely et al., 2012; Rosu et al., 2013). In the region of the germ line where DSBs are made, the X chromosome is transcriptionally quiescent and packaged in a heterochromatin-like state (Fong et al., 2002; Kelly et al., 2002). By contrast, the autosomes are transcriptionally active and replete with histone modifications associated with open chromatin. It is thought that the heterochromatic-like state of the X presents a barrier to DSB formation necessitating the evolution of proteins to overcomes this obstacle (Wagner et al., 2010). xnd-1 transcriptionally regulates him-5 to promote breaks on the X (McClendon et al., 2016), but it also restricts histone acetylation, a function that ensures the proper timing of DSB induction (Wagner et al, 2010). In the absence of xnd-1, there is precocious formation of DSBs as monitored by the accumulation of the HR strand-exchange protein RAD-51 in transition zone nuclei (Gao et al., 2015; McClendon et al., 2016; Wagner et al., 2010).

CO positioning is altered in all situations where DSB levels are affected, but the rec-1 mutation has a profound effect on the distribution of meiotic COs (Rose & Baillie, 1979) with only a minor reduction in total DSB numbers (Chung et al., 2015). Prior analysis of rec-1 and him-5 found these genes to be paralogs that function cooperatively to ensure wild-type DSB levels (Chung et al., 2015). Interestingly, while loss of rec-1 preferentially affects CO distribution, loss of him-5 reduces DSB numbers nearly in half, severely delays DSB induction (until mid-late pachytene), and changes CO distribution (Meneely et al., 2012). him-5 has redundant functions with cep-1 for both break formation and inhibition of NHEJ (Mateo et al., 2016), revealing roles for him-5 in coordinating DSB induction with downstream repair.

In this study, we expand the analysis of the genetic interactions between the SPO-11 accessory factors in C. elegans and provide the first protein-protein interaction studies of these factors using co-immunoprecipitation and yeast-two hybrid (Y2H). We use these interaction data to propose how the DSB formation is initiated in C. elegans and compare and contrast this to models proposed in yeast and plants (Bouuaert et al., 2021, Vrielynck et al., 2021). Our results allow us to begin to elucidate the regulatory events that control the localization, timing, and placement of meiotic DSBs in C. elegans.

Results

Earlier version: HIM-5 is the essential factor for meiotic breaks in the X-chromosome

Revised version: HIM-5 is a limiting factor for the HIM-17/XND-1 dependent breaks on the X-chromosome

Three of the DSB proteins appear to have a biased effect on X chromosome CO formation: null alleles of xnd-1 and him-5 predominantly show a loss of CO intermediates on the X, increases in univalent X chromosomes at diakinesis in half or almost all nuclei, respectively, and an increase in XO male offspring resulting from nondisjunction of the X chromosome during the meiotic division; the hypomorphic him-17(e2806) mutation shows increased incidence of one unattached homolog pair in diakinesis and an increase in male production consistent with an X-specific defect (Reddy & Villeneuve, 2004). To determine if these genes function through a common underlying mechanism, we first performed yeast two-hybrid (Y2H) assessing interactions with both high and low stringency. We observed a strong Y2H interaction between XND-1 and HIM-17, and a weak one between HIM-17 and HIM-5 (Figure 1A and Figure S1). To validate these interactions, we also performed immunoprecipitations (IPs) of HIM-17::GFP followed by mass spectrometry (MS, Figure S1B). This identified XND-1, as well as 16 additional strong interactors of HIM-17::GFP (Table S2). Moreover, endogenous XND-1 was identified in western blots of IPs performed independently using 3xHA-tagged HIM-17 epitopes, indicating a robust interaction between these proteins (Figure 1B and S1C). In contrast, HIM-5 was not detected in either the HIM-17::GFP co-IPs or IP/MS, suggesting that the interaction observed by Y2H either exists very transiently or was bridged by a protein expressed in yeast.

XND-1 and HIM-17 interact and are predominantly localized on DNA loops.

A) Summary of the yeast two-hybrid assay results. The complete set of results is given in Figure S1A. B) Western blot analysis of endogenous XND-1 (top) and HIM-5::GFP::FLAG (bottom) on HA pull-downs performed in him-17::3xHA; him-5::GFP::3xFLAG strain. Red box indicates the XND-1 band. Analysis was performed in biological duplicates. Specificity of the anti-XND-1 antibody is shown in Figure S3A. C) Quantification of embryonic viability (hatching rates) and frequencies of male offspring among the progeny of the indicated genotypes. Data are shown as mean +/- SD; * p<0.05, **** p < 0.0001, ND= not determined, n= number of P0 parents whose brood was examined. D) Representative confocal images of mid-pachytene stage nuclei showing that juxtaposed XND-1 and HIM-17 localization away from the chromosome axes (stained with anti-HTP-3). The X chromosome is not stained by either XND-1 and HIM-17 (Red arrows). Scale bars = 20μm. E) Quantification of the pairwise overlap between XND-1, HIM-17, and the chromosome axes. F) Alphafold3 model of XND-1-HIM-17 dimer. The black line shows the 180 Å distance between the HA tag in HIM-17 and the XND-1 protein which is recognized by the polyclonal antibody.

Prior studies from our group indicate that xnd-1 regulates HIM-5 transcriptionally since ectopic expression of him-5::GFP using the pie-1 germline promoter, but not expression using the endogenous him-5 promoter, rescues the X chromosome nondisjunction defects conferred by the xnd-1 mutation (Wagner et al., 2010, Meneely et al., 2012; McClendon et al., 2016). Since RNA-Seq studies showed that him-17 regulates many germline genes including him-5 (Carelli et al., 2022), we set out to determine if ectopic expression of him-5 could also rescue him-17 null mutants. Similar to what we observed with xnd-1 mutants, Ppie-1::him-5 — but not Phim-5::him-5 — rescued the X chromosome defects seen by the increase of XO male offspring from XX hermaphrodites (known as the high incidence of males or “Him” phenotype) (35% of males in him-17(ok424) vs. 7.02% in Ppie1::him-5;him-17(ok424)) (Figure 1C). Consistent with the lack of rescue, there was no statistically significant change in univalent formation in Phim-5::him-5;him- 17(ok424) (Figure S2). These results cannot be explained by HIM-5 protein levels which are lower in Ppie-1::him-5 than in Phim-5::him-5 (McClendon et al., 2016). The lack of rescue from Phim-5::him-5 supports the conclusion that HIM-17 directly regulates him-5 expression. Since the pie-1 promoter appears to be independent of HIM-17, the Ppie- 1::him-5 construct can—at least partially—bypass this regulation and suppress the him- 17 mutant phenotype. These findings suggest that both XND-1 and HIM-17 both promote X chromosome DSB and CO formation by regulating him-5 expression. Moreover, xnd-1 and him-17 have distinct roles in DSB/CO formation since only null alleles of him-17, but not xnd-1 or him-5, confer a severe loss of DSBs (Meneely et al., 2012; Reddy & Villeneuve, 2004; Wagner et al., 2010), while only xnd-1 mutations influence the timing of DSB induction/repair (Wagner et al., 2010; McClendon et al., 2016).

We note that ectopic HIM-5 expression strongly rescued the embryonic lethality conferred by the him-17 null allele (Figure 1C). While a small fraction of the lethality is attributed to X chromosome nondisjunction, most of the lethality is a consequence of autosomal missegregation. Thus, even though him-5 null mutant animals exhibit a profound loss of DSBs/COs on the X chromosomes with rare CO defects on autosomes (Meneely et al., 2012), him-5 expression can restore CO formation to both autosomes and X chromosomes. This strongly supports a general role for HIM-5 in CO formation and further indicates that HIM-5 is sufficient to restore COs on the X chromosome.

To further characterize the interaction between XND-1 and HIM-17, we also sought to determine whether they colocalize on meiotic chromosomes using confocal microscopy. Prior studies have shown that HIM-17 associates with meiotic chromatin (Reddy & Villeneuve, 2004). Co-staining the HIM-17::HA transgenic line with anti-HA and anti-XND-1 antibodies showed that HIM-17, like XND-1, is enriched on autosomes (Figure 1D; Wagner et al., 2010). Both proteins showed no overlap with the chromosome axes (marked with HTP-3 (Figure 1D) or DAPI (Figure S3B), although we cannot rule out that a small fraction of the population, either below the level of detection or transiently, localizes to the axes. These immunolocalization studies also show a strong juxtaposition of HIM-17 and XND-1, with no detectable overlap between their staining patterns in germ lines from the mitotic region through late pachytene, when XND-1 disappears from the nucleus (Figure 1E and Figure S3B). Structural modeling of the XND-1 and HIM-17 interaction (Abramson et al., 2024; Comeau et al, 2004) provides a potential explanation for this observation: HIM-17 is a very large protein, and the HA-tag positioned at its amino terminus is predicted to be located at least 180 Å away from the protein:protein interaction interface (Figure 1F).

DSB-1 mediates interactions with SPO-11 and promotes HIM-5 nuclear localization

Since HIM-5 appears to be central to DSB/CO formation on the X chromosome and the presence of DSB-1 and DSB-2 are required for all meiotic breaks, we next sought to determine whether HIM-5 interacts with DSB-1, its paralog DSB-2, or DSB-3. We found that DSB-1 and DSB-2 show robust Y2H interactions and that DSB-1 and DSB-3 exhibit weak Y2H interactions (Figure 2A, Figure S1A), confirming results from a prior study (Hinman et al., 2021). We also found that HIM-5 shows weak Y2H interactions with DSB-1 (Figure 2A). We confirmed the interaction between DSB-1 and HIM-5 with co-IP, pulling down GFP::DSB-1 and probing for HIM-5::3xHA by Western blotting (Figure 2B).

DSB-1 associates with HIM-5 and regulates its localization

A) Summary of yeast two-hybrid results with dsb-1, dsb-2, dsb-3 and him-5. The complete set of results is given in Figure S1A. B) Western blot analysis of anti-GFP pull-downs performed in GFP::dsb-1;him-5::3XHA strain showing co-IP of HA-tagged HIM-5 proteins. Analysis was performed in biological duplicates. C) Immunofluorescence analysis of HIM-5 localization in Phim-5::him-5::GFP;him-5 and Phim-5::him-5::GFP;dsb-1 backgrounds. DNA (DAPI, gray), HIM-5::GFP (anti-GFP, green), and chromosome axes (anti-pHTP-3, magenta). pHTP-3 staining delineates entrance to the transition zone (marked dotted yellow lines), while arrows indicate the pre-meiotic localization of HIM-5. D) Quantification of DAPI-stained bodies at diakinesis for the indicated genotypes. Colors correspond to the number of DAPI-stained bodies shown in the key below for N2, dsb-1, and dsb-2 mutants expressing Ppie- 1::him-5::GFP. Sample sizes (N) are indicated. Statistical significance for comparisons between groups is shown at the top (**p < 0.01). E) Quantification of embryonic viability (hatching rates) and frequencies of male offspring among the progeny of the indicated genotypes. Data are shown as mean +/- SD; NS stands for not significant; *** p < 0.001, n= number of P0 parents whose brood was examined.

We next wanted to understand the functional relevance of the interaction between DSB-1 and HIM-5. To this end, we examined the localization of HIM-5 proteins in the dsb-1 mutant background. Co-staining of the chromosome axis (anti-HTP-3) and either HIM-5::GFP transgene (anti-GFP) or endogenously-tagged HIM-5::3xHA (anti-HA) revealed that HIM-5 nuclear localization begins just prior to the transition zone in both wild type and dsb-1 mutants (Figure 2C and Figure S4A, arrows). However, whereas in wild type, HIM-5 protein remains nuclear throughout early-mid pachytene and become further enriched on meiotic chromatin; in dsb-1 mutants, HIM-5 nuclear localization is lost (Figures 2C and Figure S4B). From this change in HIM-5 protein distribution, we infer that DSB-1 (and/or its DSB-promoting function) is required to retain HIM-5 in meiotic nuclei.

Since HIM-5 levels could also be affected by dsb-1 and/or dsb-2, we next asked whether ectopic expression of HIM-5 could suppress defects in bivalent formation due to the loss of dsb-1 or dsb-2. As shown in Figure 2D, ectopic HIM-5 reduced the number of univalents in both mutant backgrounds. Since a small fraction of dsb-2; Ppie-1::him-5 animals exhibited 6 or 7 DAPI bodies, we assayed whether the transgene could improve hatching rates and reduce male production in dsb-2 (Figure 2E). Whereas the Ppie-1::him-5 transgene has no impact of hatching and male frequencies, the more highly expressed Phim-5::him-5 transgene (McClendon et al., 2016) conferred a mild suppression of the embryonic lethality. This ability to suppress hatching-but not male production-further supports a role for HIM-5 in regulating total DSB numbers.

Genetic interactions define four functional groups that contribute to DSB/CO formation

Having established interactions between subsets of the SPO-11 accessory factors in C. elegans, we next wanted to perform more extensive genetic analyses between the known DSB factors. To this end, we created double mutant combinations of alleles with partial DSB defects. These included either weak loss-of-function alleles of essential DSB genes (him-17, mre-11) or null alleles of genes that partially reduce DSBs (cep-1, dsb-2, rec-1, him-5, and parg-1). Diakinesis oocytes were analyzed in the single and double mutants by whole-mount fixation and DAPI staining followed by confocal microscopy and 3D visualization. The chromosomes are highly condensed at diakinesis, and six bivalents can be detected in wild-type worms. Univalents, fusions, and DNA fragments can be seen in mutants with defects in DSB formation or their repair into a CO, respectively (Chin & Villeneuve, 2001; Guo et al., 2022; Rosu et al., 2013). We hypothesized that double mutant strains carrying mutations in genes that function in different aspects of DSB induction would show more severe defects and have an increase in univalent formation compared to single mutants. On the other hand, genes that function together or at the same step of break induction should exhibit no or only mild enhancement of the DSB defects.

Double mutant analysis was precluded with dsb-1 and dsb-3 since the single mutants already show a complete absence of breaks (Hinman et al., 2021; Stamper et al., 2013). By contrast, dsb-2 mutant nuclei manifest an age-dependent loss in DSB capacity (Rosu et al., 2013; Stamper et al., 2013) allowing us to analyze interactions between dsb-2 and cep-1, parg-1, him-17, and rec-1 (Figure 3A and Figure S5A-D). As 1-day-old adults, dsb-2 single mutants had ∼50% of diakinesis nuclei that contained 6 bivalents, ∼30% with 7 DAPI-stained bodies (5 bivalents and 2 univalents), and ∼20% with 8 or more DAPI-stained bodies. Single mutants of cep-1, parg-1, and rec-1 each had >90% of diakinesis nuclei with 6 bivalents, consistent with prior studies (Chung et al., 2015; Janisiw et al., 2020; Mateo et al., 2016). In animals carrying him-17(e2806) mutation, a weak allele of the otherwise essential locus for DSB formation, only ∼50% of mutant nuclei had 6 bivalents; the other 50% had 7 DAPI-stained bodies, reflecting the requirement for this gene in CO formation on the X chromosome (Reddy & Villeneuve, 2004). When dsb-2 was combined with each of these mutations, an increased number of DAPI-stained bodies and a concomitant decrease in the number of chiasmata were observed (Figure 3A, Figure S5). These results strongly suggest that dsb-2 functions cooperatively with all four loci for the regulation of meiotic CO formation.

Genetic and protein-protein interaction studies show evidence of SPO-11 accessory protein sub-complexes in C. elegans.

A) Summary of the genetic interaction results. The complete set of results is given in Figure S5. B) Genetic groupings based on data presented here and previously published (Chung et al., 2015; Mateo et al., 2016; McClendon et al., 2016; Janisiw et al., 2020). C) Summary of yeast two-hybrid results. The complete set of results are shown in Figure S1A. D) Schematic representation of the interaction network based on Y2H interactions. Thick line: Strong interaction; thin line: weak interaction. E, F) Yeast-two-hybrid interactions of SPO-11, DSB-2, and DSB-3 with different DSB-1 sub-domains (schematic in F). Details on DSB-1 sequence, protein structure, and deletions used in these experiments, are found in Figure S8. G) AlphaFold2 predictions of DSB-1, DSB-2 and DSB-3 protein structures. H, I) Predicted interaction domains of DSB-1, DSB-2 with DSB-3 based on homology with the yeast proteins. The DSB-1/-1/-3 trimer is thermodynamically more stable than the DSB-1/-2/-3 trimer, consistent with the subordinate role for DSB-2 in break induction in young animals.

Given that rec-1, parg-1, and cep-1 fall in a separate functional group from dsb-2, we next addressed whether these three genes function together or independently. All double mutant permutations of rec-1, parg-1, and cep-1 were analyzed, and in each case, the single and double mutants were nearly indistinguishable from the single mutants, i.e., 6 bivalents (Figure S5L-M). Therefore, we infer that these genes likely function to control the same step in DSB/CO induction (Figure 3B). Consistent with this interpretation, when rec-1, parg-1, or cep-1 were combined with other DSB gene mutations, including him-17, they had similar effects, and the numbers of univalents were increased (Figures S54A, C, D for Group 1; Figure S5E-G for Group 3; and Figure S5J, K for Group 4). Similar results have previously been reported from our laboratories for double mutants with him-5 (Chung et al., 2015; Janisiw et al., 2020; Mateo et al., 2016).

mre-11 functions in both DSB formation and subsequent DNA end resection (Chin & Villeneuve, 2001), the latter as part of the MRE11-RAD50-NBS1/Xrs2 (MRN/X) resection complex, the other proteins of which are encoded by the worm rad-50 and nbs-1 genes (Girard et al., 2018; Hayashi et al., 2007). Null mutations in mre-11 and rad-50 completely abrogate DSB induction. However, we were able to take advantage of a separation-of-function allele, mre-11(iow1), that cannot perform its end-resection functions but can induce DSBs (Yin & Smolikove, 2013). The iow1 mutation causes chromosome fusions to appear in diakinesis oocytes resulting from the aggregation of unprocessed DNA ends. While this allele makes enough DSBs to cause aggregation of all chromosomes, we hypothesized that it might have a minor impairment in DSB induction that would be revealed when combined with other DSB mutations. If this were the case, the double mutants would make fewer breaks, leading to fewer chromosome fusions and an increase in the number of DAPI-stained bodies/univalents. Consistent with this hypothesis, mre-11(iow1) double mutants with each of the other DSB mutations led to fewer fusions and more DAPI-stained bodies (Figure 3A, Figure S5H-K). Thus, we propose that mre-11, and by extension rad-50, function in their own branch of the DSB/CO pathway.

Surprisingly, a unique phenotype was observed in a subset of the cep-1; mre-11 double mutants. Our prior studies showed that cep-1 functions redundantly with him-5 to both ensure DSB formation and promote downstream homologous recombination (HR) repair by preventing non-homologous end joining (NHEJ) (Mateo et al., 2016). In that study, we reported that cep-1(lg12501) is a separation-of-function allele that is defective in DSB formation. In the cep-1(lg12501); mre-11(iow-1) double mutants, we observed mixed phenotypes: some nuclei had 12 DAPI-stained bodies, while others exhibited an exacerbation of the fusion defect (Figure S6). We infer that the univalents indicate a synergistic effect between mre-11 and cep-1 for DSB formation; the increase in fusions, on the other hand, may suggest that cep-1(lg12501) is not fully wild-type for the DNA repair functions of CEP-1.

While each of the genes that we analyzed has reported roles in DSB formation, we wanted to confirm that the observed increases in univalents in the double mutants were in fact due to impairment in DSB induction and not to deficits in later steps of CO formation. Therefore, we relied on an established assay to distinguish between these possibilities: we asked whether the addition of exogenous breaks from gamma irradiation (IR) could suppress the univalent phenotype by serving as a surrogate for SPO-11-induced DSBs (Dernburg et al., 1998). Recent work has shown that 10Gy IR is able to produce up to ∼20 DSBs per meiotic nucleus (Lascarez-Lagunas et al., 2023), which is sufficient to ensure a CO on each chromosome (Mets & Meyer, 2009). We analyzed the number of DAPI bodies with and without IR in a subset of the double mutants where we observed synergistic effects (dsb-2; parg-1, dsb-2; rec-1, and rec-1; him-17). When each of these double mutant strains was exposed to IR, the wild-type phenotype of 6 bivalents was largely restored (Figure S7), supporting the conclusion that the CO deficiency observed is triggered by defects in DSB formation.

Protein-protein interactions between the DSB regulatory factors

The genetic analyses described above defined four functional groups that contribute to break formation (Figure 3B). Using Y2H, we next set out to determine whether interactions between and within groups could explain how the different complexes come together to regulate DSB induction (Figure 3C-D, Figure S1A). CEP-1 showed auto-activation in both vectors and therefore was removed from further analysis. DSB-1, HIM-17, and REC-1 also showed self-activation in the binding domain vector, therefore, only the activation domain fusions were analyzed. Pairwise interactions between all remaining combinations were then analyzed. HIM-5 displayed the most promiscuous Y2H interactions, with strong interactions seen with MRE-11 and weak interactions as already discussed with HIM-17 and DSB-1. XND-1 apart from interacting with HIM-17, also showed strong Y2H interactions with REC-1. A recent study provided Y2H evidence supporting the idea that SPO-11 interacts with DSB-1 (Hinman et al., 2021). Critically, our work supports and extends this observation to show that DSB-1 is the only known accessory factor that interacts with SPO-11 in Y2H assays, cementing its role as a key determinant for DSB induction and explaining its role in feedback controls (Stamper et al., 2013; Hinman et al., 2021; Guo et al., 2022).

To analyze more in-depth the Y2H interactions between DSB-1 and DSB-2, DSB-3, and SPO-11, we looked for protein domains in DSB-1 that could be involved in these interactions. Interestingly, DSB-1 has 5 “sticky” helices (helices 2-6 in Figure S8) that have the potential to be implicated in protein-protein interactions. We first deleted all of these helices together by introducing a deletion after L114 (DSB-1(1-114) in Figures 3E, F). We found that this truncated version of DSB-1 loses the ability to interact with DSB-2 and DSB-3 but is still able to associate with SPO-11. Moreover, when we only deleted the helices 3-6 but kept the helix 2 (DSB-1(1-145)), we obtained the same results, suggesting that the Y2H interactions between DSB-1 and DSB-2/DSB-3 is mediated by the C-terminal helices 3-6 and not helix 2. To prove this, we created a truncated DSB-1 protein that only contains amino acids from P145 to the stop codon (DSB-1(145-Ct)). Consistent with the deletion analysis, we observed that the C-terminal domain is sufficient to promote associations with DSB-2 and DSB-3. This domain was also not sufficient to promote the DSB-1/SPO-11 interaction (Figure 3E). These experiments complement the results published by Villeneuve lab where they observed that the DSB-1 N-terminus is the domain necessary for the DSB-1/SPO-11 (but not for the DSB-1/2/3) interaction (Hinman et al., 2021).

We next used AlphaFold and structural modeling to investigate whether the C. elegans REC114/MEI4 homolog could assemble in the same fashion as their distant orthologues (Claeys Bouuaert et al., 2021; Daccache et al., 2023) in which two helices of REC114 dimer surround a MEI4 helix. DSB-1 is predicted to form a stable homodimer that can assemble into a highly stable trimeric complex with DSB-3 threaded through the middle (Figure 3G-I) analogous to what is predicted for Rec114-Mei4 (Daccache et al., 2023). By contrast, DSB-2 cannot form homodimers, although it can replace one of the DSB-1 monomers to form the DSB-1:DSB-2:DSB-3 ternary complex. Free energy calculations (Champ and Camacho, 2007) show that the DSB-1:DSB-2:DSB-3 complex is less stable (−3.8 kcal/mol) than a DSB-1:DSB-1:DSB-3 complex consistent with the more limited role for DSB-2 in break formation, at least in young worms (Rosu et al., 2013; Stamper et al., 2013)).

Discussion

Earlier version: Across species, accessory factors collaborate with SPO-11 to regulate various aspects of DSB formation, including the timing, placement, number, and persistence of DSB induction, as well as the chromatin architecture at and around the break site (Arora et al., 2004; Cole et al., 2010; Hinman et al., 2021; Kumar & de Massy, 2010; Panizza et al., 2011; Sasanuma et al., 2007; Sommermeyer et al., 2013; Vrielynck et al., 2021). Here we present genetic, biochemical, and cytological studies to develop a working interaction model for DSB regulation in C. elegans. The C. elegans DSB complex appears to be comprised of four sub-groups (Figure 3B and Figure 4) interconnected by their association with HIM-5. Our results suggest two important roles for HIM-5: first as the linchpin between the whole DSB complex, thereby ensuring the coordination of the timing, placement, and number of DSBs (described in more detail below); and second, as the determinant in X chromosome COs.

Revised version: Across species, accessory factors collaborate with SPO-11 to regulate various aspects of DSB formation, including the timing, placement, number, and persistence of DSB induction, as well as the chromatin architecture at and around the break site (Arora et al., 2004; Cole et al., 2010; Hinman et al., 2021; Kumar et al., 2010b; Panizza et al., 2011; Sasanuma et al., 2007; Sommermeyer et al., 2013; Vrielynck et al., 2021). Here we present genetic, biochemical, and cytological studies to develop a working interaction model for DSB regulation in C. elegans. The C. elegans DSB complex appears to be comprised of four sub-groups (Figure 3B and Figure 4) interconnected by their association with HIM-5. Our results suggest two important roles for HIM-5: first as the linchpin between the whole DSB complex, thereby ensuring the coordination of the timing, placement, and number of DSBs (described in more detail below); and second, as the determinant in X chromosome COs.

Speculative model of the DSB formation complex in C. elegans

Synthesizing prior work and data herein, we propose a model to explain the known interactions and localization patterns of the DSB regulatory factors in C. elegans. DSB-1/-2/-3 and HIM-17-XND-1 subcomplexes are initially found on DNA loops. The phosphorylation of REC-1 by CDK activates DSB formation, allowing HIM-5 through its association with MRE-11 and PARG-1 (which both associate with HTP-3) to efficiently bring the DSB-1/-2/-3 complex with SPO-11 to the chromosome axis. This recruitment to the axes by HIM-5, facilitates the coupling of DSBs to downstream repair by the MRN complex and others. In the absence of HIM-5, DSBs/COs on autosomes do not occur efficiently and crossover sites are shifted to the gene-rich third of each chromosome, presumably because DSB-1/-2/-3/SPO-11 cannot be as efficiently recruited to the proper sites partially defined by HIM-17 and XND-1. We assume that recruitment to the chromosome axes in the absence of HIM-5 is likely stochastic or mediated by as yet unknown protein interactions. On the X chromosome, the absence of HIM-5 prevents the formation of most DSBs.

DSB-1 and DSB-2 are both orthologs of the budding yeast Rec114 protein but DSB-1 has a more central role in DSB induction (Rosu et al., 2013; Stamper et al., 2013) which we posit can be explained by the increased stability of the DSB-1:DSB-1:DSB-3 trimer compared to a DSB-1:DSB-2:DSB-3 trimer. The direct association between SPO-11 and DSB-1 but none of the other worm accessory factors points to DSB-1/REC114 as the hub for the integration of break inducing signals in this organism. In yeast and Arabidopsis, accessory proteins are associated indirectly with Spo11/SPO11 through interactions with the TopVI/MTOPVIB subunit (Arora et al., 2004; Kee & Keeney, 2002; Maleki et al., 2007; Robert et al., 2016; Vrielynck et al., 2021). These differences between species underscore how conserved DSB proteins have been co-opted to regulate different aspects of break induction.

While DSB-1/-2/-3 are critical for CO formation, they predominantly localize on DNA loops with little to no overlap (Hinman et al., 2021). Instead, a small fraction of these three proteins appear to associate with the DNA axis, a localization that is consistent with the “tethered loop model” for break induction in yeast and plants (Kee et al., 2004; Maleki et al., 2007; Panizza et al., 2011; Sommermeyer et al., 2013; Tsai et al., 2020). Two accessory factors, MRE-11 and PARG-1 are known interactors of the axis protein HTP-3 (Janisiw et al., 2020). Both proteins interact either directly or indirectly with HIM-5, which physically associates with DSB-1, thereby providing a mechanism to bring the loop-associated proteins to the axis (Figure 4).

MRE-11 and PARG-1 also have roles in both break formation and DNA repair. For PARG-1, we show here that it functions together with REC-1 and CEP-1 based on the following data: synthetic CO defects were seen in cep-1;him-5, rec-1;him-5, and parg-1;him-5 (Chung et al., 2015; Mateo et al., 2016 and this paper), but not in rec-1;cep-1, rec-1;parg-1, and cep-1;parg-1 double mutants; dsb-2 mutant phenotypes are exacerbated by the loss of rec-1, cep-1 (Figure S5), and parg-1 (Janisiw et al., 2020). Interestingly, both CEP-1 and PARG-1 have additional roles in promoting homologous recombination-mediated repair and/or preventing alternative repair pathways (Janisiw, et al., 2020; Mateo et al., 2016). We envision that CEP-1, REC-1, and PARG-1 function together to promote DSBs and their repair into COs at the chromosome axis. Since we also showed that the DSB defects in mre-11 mutant animals were enhanced by loss of parg-1 or rec-1 (Figure S5), we favor a model in which MRE-11 and PARG-1 independently interact with the chromosome axis to ensure robust DSB formation and/or repair. It is possible that MRE-11, which is essential for COs, facilitates the tethering of chromosome loops to the axis, but further structural studies of these complexes are required.

The defects observed in cep-1;him-5 (Mateo et al., 2016), and rec-1;him-5 (Chung et al., 2015) are both exacerbated by maternal age, similar to dsb-2. We envision that inputs from both HIM-5 and DSB-2 converge on DSB-1 to ensure robust DSB formation across the lifespan. This would also explain the synthetic defects in DSB formation between dsb-2 and each of the components in the him-5 interaction network. DSB-1 is required for chromatin association of DSB-2 but not vice versa, although DSB-1 levels are reduced in older dsb-2 mutant animals (Rosu et al., 2013; Stamper et al., 2013). We show here that DSB-1 is also required for nuclear retention of HIM-5. While the functional significance of this interaction awaits isolation of separation-of-function mutations that abrogate their interaction, it is tempting to speculate the signals that initiate meiotic induction regulate this interaction.

TIMING of BREAKS

Studies from multiple systems have elucidated mechanisms to coordinate the timing of DSBs with the completion of meiosis, e.g. by CDK and Ddf4 phosphorylation in yeast (Murakami & Keeney, 2008; Carelli et al., 2022). In worms, REC-1 was previously shown to be a target of CDK and therefore it was hypothesized that such a phosphorylation event would be critical to ensure that break induction occurs after the completion of S phase (Chung et al., 2015). REC-1 and HIM-5 evolved from a gene duplication event in the C. elegans lineage; whereas the ancestral gene in related nematodes has features of both proteins-although it shares greater homology with HIM-5 (Chung et al., 2015). Thus, we hypothesize that the single ancestral REC-1/HIM-5 protein would be subject to CDK phosphorylation, which would then promote the assembly of the complexes at the chromosome axis to promote break formation.

Timing also appears to be critical for CO outcomes. Studies in yeast show different outcomes for early and late breaks (Joshi et al., 2015) and recent studies in worms indicate that later breaks are needed for CO formation (Hicks et al., 2022) . We previously showed that xnd-1 mutants have altered DSB kinetics, showing elevated levels of DSB formation in leptotene/zygotene (McClendon et al., 2016; Wagner et al., 2010). This function of XND-1 is mediated, at least in part, by its effect on chromatin architecture. One possibility is that the physical interaction that we show here between XND-1 and REC-1 couples the completion of S phase to the activation of breaks within the correct chromatin environment, thus ensuring their timely formation. In the absence of XND-1, increases in global histone acetylation (Gao et al., 2015; Wagner et al., 2010) might allow SPO-11 easier access to the DNA. Another possibility is that upon phosphorylation of REC-1 by CDK-1, REC-1 through PARG-1 transiently brings XND-1 (the timer) to the axis, an event that restrains DSB formation.

Chromatin and Meiotic Gene Transcription

We provide strong evidence for direct interactions between XND-1 and HIM-17, both by Y2H and co-IP. HIM-17 appears to have a role in transcriptional regulation of DSB factors. Here we show that ectopic expression of HIM-5 from the pie-1 promoter, but not overexpression from the endogenous him-5 promoter, suppressed him-17 phenotypes. This is similar to the rescue we saw of xnd-1 (McClendon et al., 2016) and suggests that both proteins contribute to the expression of him-5. This supports findings that him-17 regulates the expression of ∼300 germline-enriched genes, including him-5, rec-1, and dsb-2, (Carelli et al., 2022) whose concomitant loss would be expected to abrogate break induction (Chung et al., 2015 and Figure 3A and Figure S5). While its only role may be in transcriptional regulation, a growing literature and data presented herein suggest a more direct role for HIM-17 in meiosis. We presented Y2H data for a weak interaction between HIM-17 and HIM-5. Although these interactions were not confirmed by IP, it remains possible that these complexes are transient associations, are unstable during IP, and/or are not abundant. Together with observed changes in chromatin structure and both DSB and CO distribution in him-17 mutants (Nadarajan et al., 2021; Reddy & Villeneuve, 2004), it is tempting to speculate that together with XND-1, HIM-17 helps to create an environment permissive for SPO-11 binding and/or break formation. XND-1 and HIM-17 may function together in this role or redundantly, perhaps through XND-1 interacting with REC-1 and HIM-17 interacting with HIM-5, or in related species through interactions with the single REC-1/HIM-5 paralog.

X/Autosome differences in DSB formation

One of the conundrums arising from our data is in the function of the HIM-5 protein. him-5 was identified based on its preferential impact on CO formation on the X chromosome (Broverman & Meneely, 1994; Hodgkin et al., 1979). The X-specific phenotypes conferred by him-5 mutation could suggest that the X chromosome is simply more sensitive to a reduction in the number of active SPO-11 complexes; the fact that mutations in dsb-2 also reduce DSBs but without an X chromosome bias argues against this interpretation. We show herein that HIM-5 is the sufficient to ensure that COs occur on the X chromosome. Ectopic expression of him-5 can substantially rescue both the autosomal and X chromosomal defects of him-17 null mutations -- despite expression levels lower than endogenous him-5 (McClendon et al., 2016). The X chromosome-specific defects of him-5 mutations might arise due to defects in recruitment of DSB-1-2-3/SPO-11 to specific sites on the X or to an impaired ability to coordinate the assembly of the DSB machinery on the— late-replicating (Jaramillo-Lambert et al, 2007)— axis of the X chromosome, although other models can be envisioned.

Our data and a growing literature insinuate a more general and central role for HIM-5 in DSB formation and downstream CO repair: 1) rec-1;him-5 double mutants give an age-dependent severe loss of DSBs (like dsb-2 mutants) suggesting that the ancestral function of the protein may have a more profound effect on break formation (Chung et al 2015); 2) total DSBs are reduced to nearly 50% of wild type and are severely delayed in formation in him-5 mutants (Meneely et al., 2012); 3) a recent study reports role for him-5 in CO formation in males where the X chromosome does not engage in CO exchange (Engebrecht et al., 2025); 4) HIM-5 is required for the normal distribution of COs on the autosomes (Meneely et al., 2012); and 5) redundantly with CEP-1, HIM-5 promotes DSB formation and HR-mediate repair and inhibits NHEJ (Mateo et al, 2016), 5) Our data herein show that HIM-5 interacts not only with DSB-1, but also components of the other functional groups that regulate DSB induction. Together, we believe this data supports a role for HIM-5 in bridging the accessory factors with SPO-11 to ensure the proper timing, placement and number of breaks.

Concluding Remarks

Earlier version: Our Y2H and IPs results show very strong protein-protein interactions between the different DSB factors, however, their localization along the gonads does not frequentlyt= coincide suggesting temporal and spatial regulation throughout meiosis. Further studies are necessary to elucidate how these interactions arise and change in vivo and how these are integrated with both cell cycle and regulatory feedback controls.

Revised version: Our Y2H and IPs results show very strong protein-protein interactions between the different DSB factors, however, their localization along the gonads does not frequently= coincide suggesting temporal and spatial regulation throughout meiosis. Further studies are necessary to elucidate how these interactions arise and change in vivo and how these are integrated with both cell cycle and regulatory feedback controls.

Methods

Culture and strains

Worms were cultured on MyoB plates (Burns et al., 2006)) seeded with OP50 at 20°C, unless otherwise noted (Brenner 1974). Mutant strains used in this study are listed in Supplemental Material, Table S1. All strains were derived from the wild-type Bristol strain N2.

Chromosome morphology analysis of diakinesis oocytes

The numbers of DAPI-stained bodies present in diakinesis oocytes were assessed in intact adult hermaphrodites at day 1 of adulthood (24 hours post-L4, unless otherwise specified). Adults were fixed in Carnoy’s fixative solution (60% ethanol, 30% chloroform, and 10% glacial acetic acid) and stained with 4′,6-diamidino-2-phenylindole (DAPI) (10 mg/ml stock diluted 1:50,000 in 1X Phosphate-buffered saline; PBS) for 15 minutes in a humid chamber. Worms were mounted in Prolong Gold with DAPI, cured overnight at room temperature, and stored at 4 °C prior to imaging either on a Nikon A1r confocal microscope or Leica Stellaris 5 confocal with an integrated White Light Laser (WLL). Diakinesis images were procured as 0.2µm per plain Z-stacks and visualized using Volocity 3-D Imaging Software for Nikon images and LAS X for Leica images. Fisher tests were performed to analyze the DAPI-stained chromosome morphology data in diakinesis oocytes. This approach was used to compare the distribution of DAPI-stained bodies, focusing on differences in the frequency of univalents and bivalents. Statistical analyses were conducted using R software, with a significance threshold set at p < 0.05.

γ-irradiation

Worms of specified genotypes were exposed to 10Gy of γ-irradiation using a 137Cs source (Gammacell1000 Elite; Nordion International). Analysis of diakinesis oocytes by DAPI staining was performed as described above at 28 h post-irradiation.

Immunostaining

Gonads from day 1 adults of the appropriate genotypes were dissected in 1X sperm salts (50 mM PIPES pH 7.0, 25 mM KCl, 1 mM MgSO4, 45 mM NaCl, and 2 mM CaCl2) with 1 mM levamisole and fixed in 1% paraformaldehyde diluted in 1X sperm salts for 5 min in a humid chamber. Slides were then frozen on a metal block on dry ice for at least 10 min prior to flicking off the cover slip and immersing in 100% ethanol at -20° C for 2 min. Slides were then washed 3X for 10 min each in PBSTB [1XPBS with 0.1% Tween and 0.1% bovine serum albumin (BSA)], and prior to overnight incubation with primary antibody also diluted in PBSTB. Primary antibodies were as follows: mouse anti-GFP 1:500 (Invitrogen Cat. #33-25—Lot 683990A), rabbit anti-phosphoHTP-3S285 1:2000 (Das et al., 2022), mouse anti-HA 1:1000 (mouse anti-HA-tag (6E2) Cell Signaling Cat. #23675, Lot 5); guinea pig anti-XND-1, 1:10,000 (Wagner et al., 2010). Overnight incubations were performed at 4° C except for anti-HA which was performed at room temperature. The next day, slides were washed 3X in PBSTB for 10 min each and incubated with Alexa-conjugated secondary antibodies from Molecular Probes: anti-mouse Alexa 488, anti-rabbit Alexa 633, anti-guinea pig Alexa 568 (all diluted 1:2000 in PBST/ BSA 0.1%). Incubation was for 2 h at room temperature in the dark. Slides were then washed 2 X 10 min in PBSTB, and 1 X 10 min with DAPI (10 mg/ml stock diluted 1:50,000 in 1X PBS). Slides were mounted in Prolong Diamond with DAPI and put in the dark to cure overnight before imaging. Images were acquired on a Leica Stellaris 5 confocal and Lightning image processing during acquisition. Z-stacks were acquired with 0.2 micron or smaller step size and visualized in 3D with Leica LAS X software.

Yeast two-hybrid assay

cDNAs were amplified from germline-specific cDNA libraries and were cloned by Gibson Assembly into the GAL4 activating domain (pGAD-C1) and binding domain (pGBD-C1)– expressing vectors. Plasmids with the desired cDNA were verified by Sanger sequencing prior to use. Plasmids were co-transformed into the PJ69-4a yeast strain and selected by growth on SC-LEU-TRP medium. Three to five transformants were grown to early log phase (OD₆₀₀ = 0.2), and 5 µL of culture was spotted onto selection plates:

SC-LEU-TRP (loading control, selecting for plasmid presence)

SC-LEU-TRP-HIS (reporter HIS3 expression, indicating interaction)

SC-LEU-TRP-ADE (more stringent selection for strong interactions)

Plates were incubated at 30°C for 72 hours and then imaged. The experiment was performed in triplicate, and empty vectors in the activating and binding domains were included as negative controls. Data are presented as spots on a plate and are from three to four individual transformants per interaction tested and are not individual colonies. The experiment was repeated in triplicate from different transformations.

Co-Immunoprecipitation and Western Blot

In order to perform co-immunoprecipitation experiments, synchronized GFP::dsb-1; him-5::3xHA and him-5::3xHA or him-17::3xHA;Phim-5::him-5::GFP::3xFLAG and Phim-1::him-5::GFP::3xFLAG strains were grown until young adult stage (24h post-L4) and nuclear protein fractionation was executed as previously shown (Silva et al.; 2014). The chromatin-bound fraction was generated upon Benzonase digestion (25U/100 μl of extract for 1h at 4°C). 2 mg of nuclear extract (nuclear-soluble and chromatin-bound fractions were pooled together) was incubated with Agarose GFP traps (Chromotek) or anti-HA Affinity Matrix (Roche) in Buffer D (20mM HEPES pH 7.9, 150mM KCl, 20% glycerol, 0.2mM EDTA, 0.2% Triton X-100 and 1x complete Roche inhibitor) overnight at 4°C. The following day, the agarose beads were recovered by centrifugation at 7500 rpm for 2’at 4°C and extensively washed with Buffer D at room temperature. After the final wash, the beads were resuspended in 40 μl of 2x Laemmli Buffer and boiled for 10’, after which they were spun at maximum speed for 1’ and the whole eluate was loaded on a precast 4-20% gradient acrylamide gel. Proteins were transferred on a nitrocellulose membrane for 90 minutes at 4°C at 100V and blocked for 1h at room temperature in 1x TBST containing 5% BSA. Mouse monoclonal anti-HA (Cell Signaling), chicken polyclonal anti-GFP (AbCam), mouse monoclonal anti-FLAG HRP-conjugated (Sigma), and polyclonal anti-XND-1 (Wagner et al., 2010) antibodies were diluted in blocking solution at 1:1000, 1:5000, 1:2000 and 1:2500 respectively and left to incubate overnight at 4°C. Washes were performed in 1x TBST and anti-mouse, anti-chicken, and anti-guinea pig HRP-conjugated secondary antibodies (ThermoFisher) were diluted 1:10,000 in 1x TBST containing 5% milk for 1h at room temperature. After several washes in 1xTBST, the membrane was incubated with Clarity Max ECL (BioRad) and imaged with a G-Box (Syngene).

Immunoprecipitation and Mass spectrometry analysis of HIM-17::GFP

IPs were performed on stock AV280: unc-119(e2498) III; him-17(ok424) V; meIs5[him-17::GFP +unc-119(+)]. Preparation and freezing of synchronized worms were carried out as in Cheeseman et al., (2004) except that worms were grown on 8P plates seeded with the E. coli strain NA22 at room temperature. Worm lysis was carried out by grinding worm pellets in liquid nitrogen and subsequent sonication (Diagenode) was performed with the following settings: 3x 5 cycles on high with 30 seconds on and 30 seconds off. Lysate was centrifuged for 45 minutes at 15000 rpm at 4°C. To avoid unspecific protein binding the worm lysate was pre-cleared by incubation to uncoupled agarose beads (Chromotek, bab-20) for 1 hour at 4°C. Pre-cleared worm lysate was incubated with GFP-trap (Chromotek, gta-20) for 3 hours at 4°C. Beads were washed 3 times with lysis buffer containing proteinase inhibitors (Roche) and 7 times with lysis buffer without proteinase inhibitors. Proteins were eluted by adding 40 μl 2xLSB and boiling for 5 min at 95°C. Protein precipitation was carried out using ProteoExtract Protein Precipitation Kit (EMD Millipore, 539180-1KIT) according to the manufacturer’s protocol and submitted for MS analysis to the Taplin Biological Mass Spectrometry Facility. The mass spectrometry proteomics data have been deposited to the ProteomeXchange Consortium via the PRIDE partner repository (Perez-Riverol et al., 2022) with the dataset identifier PXD057629 and 10.6019/PXD057629.

GFP pulldowns of HIM-17::GFP for IP/MS was performed in duplicate and corrected against unspecific binding of proteins bound to GFP trap beads. Specifically, potential false positive hits, defined as such because they were also detected in wild type (N2), AIR-2::GFP and AIR-2KD::GFP pulldowns, were removed from this list. Then, only the remaining hits detected on both independent biological repeats, and for which 2 or more peptides were detected in each experiment, made it to this final table. The only exception is KU-80 which made it to the list after subtraction of the unspecific/false positive hits, but for which only 1 peptide was detected in the December pulldown (14 peptides were detected in the Feb 2014 pulldown). Given that we found KU-70 both times and the coverage/depth was better for the Feb 2014 pulldown we decided to keep KU-80 on the final table. Supplemental Table 2 shows all the hits that remained after subtraction of the unspecific binding proteins.

Protein structure predictions

Structure analysis of protein-protein interaction

Prediction of protein-protein interactions is based on available protein structures or AlphaFold2 (AF2) (Jumper et al., 2021) models. Most relevant proteins DSB-1, DSB-2 and DSB-3 in C. Elegans do not have crystal structures. Hence, we rely on AF2 homology models that predicts core domains with confidence factors better than 90, well within the range of applicability to scan for possible protein-protein interactions. Docked poses among these proteins will be predicted using ClusPro (Comeau et al., 2004b, 2004a; Kozakov et al., 2017) and free energy complementarity will further analyze with FastContact (Camacho & Zhang, 2005; Champ and Camacho, 2007), both well-known structural modeling tools developed by one of us.

Peptide design

Peptides were designed using MD simulations in AMBER18 on the GPU-accelerated code with the AMBER ff14SB force field (Salomon-Ferrer et al., 2013; Maier et al., 2015). The tLeap binary was used to solve structures in an octahedral TIP3P water box with a 15 Å distance from the peptide surface to the box edges and a closeness parameter of 0.75 Å. The system was neutralized and solvated in 150 mM NaCl. The nonbonded interaction cutoff was set to 8 Å. Hydrogen bonds were constrained using the SHAKE algorithm and an integration time step of 2 fs. Simulations were carried out by equilibrating the system for 5 ns at NPT, using a Berendsen thermostat to maintain a constant pressure of 1 atm followed by 300 ns NVT production at 300 K.

Structural Modeling of Full-Length DSB-1, DSB-2, and DSB-3 Interactions

To model the interaction between the full-length DSB-1, DSB-2, and DSB-3 proteins, we used the Mol* Viewer (Molstar), an interactive web-based molecular visualization platform. Predicted structures were imported into the viewer to evaluate their relative orientation, potential interaction surfaces, and spatial compatibility. Full-length models were analyzed using chain-specific coloring and representation tools to facilitate comparative visualization and interpretation of possible interaction interfaces.